## `lab02`—Special Functions

![](./img/header-bkgd.png)

❖ Objectives

-   Write a simple function to implement a mathematical formula.  <!-- connor -->
-   Understand the difference between _returning_ a value and _printing_ a value.  <!-- dinlor -->
-   Use default values in functions.  <!-- luxget -->

[Special functions](https://en.wikipedia.org/wiki/Special_functions) are a wide-ranging class of named useful functions.  These include common ones like sine, arctangent, and hyperbolic cosine, as well as rare ones like Bessel and Airy functions, integral functions, and hypergeometric functions.  Now that we have learned how to define many basic processes using functions and loops, we can implement our own versions of these functions and understand how library functions may be implemented.

#### Bessel Functions

[Bessel functions](https://en.wikipedia.org/wiki/Bessel_function) describe solutions to a particular differential equation, Bessel's equation, which describes situations such as a swinging rope, vibration in a drum, diffraction from a helix, etc.

$$
x^2 \frac{d^2 y}{dx^2} + x \frac{dy}{dx} + \left(x^2 - \alpha^2 \right)y
=
0
$$

The solutions of this equation are known as $J_n$ and $Y_n$.  The first solution, $J_n$, looks like this:

$$
J_n(x)
=
\sum_{m=0}^\infty \frac{(-1)^m}{m! \Gamma(m+n+1)} {\left(\frac{x}{2}\right)}^{2m+n}
$$

![](https://upload.wikimedia.org/wikipedia/commons/5/5d/Bessel_Functions_%281st_Kind%2C_n%3D0%2C1%2C2%29.svg)

Let's implement $J_n$ using a `for` loop and the built-in `math.factorial` and `math.gamma` functions.

-   Compose a function `besselj` which accepts two values, `n` and `x`.  `besselj` should return the solution at that point according to the formula for $J_n$ above.

<!-- grade_id: besselj, points: 0, grade: false, solution: true -->
```python
# define your function here
from math import factorial, gamma
def besselj( n,x ):
    sum = 0
    for m in range(20):  # we use a large number but not actually infinity
        numer = _____
        denom = factorial(m) * gamma(m+n+1)
        term  = (x/2)**(2*m+n)
        sum = sum + numer/denom*term
    return sum
```

<!-- grade_id: besselj-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for case of m = 0
from scipy.special import j0
from numpy import isclose
assert isclose(besselj(0,1),j0(1))
assert isclose(besselj(0,2),j0(2))
print( 'Success!' )
```

<!-- grade_id: besselj-test-case, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
# test for case of m = 1
from scipy.special import j1
from numpy import isclose
assert isclose(besselj(1,1),j1(1))
assert isclose(besselj(1,2),j1(2))
print( 'Success!' )
```

#### Gamma Function

The $\Gamma$ function is a generalization of the factorial.  You just saw an implementation using `math`; let's build our own.

$$
\Gamma(z)
=
\frac{1}{z}
\prod_{n=1}^{\infty}
\left[ \frac{1}{1+\frac{z}{n}} \left(1 + \frac{1}{n}\right)^z \right]
$$

![](https://upload.wikimedia.org/wikipedia/commons/5/52/Gamma_plot.svg)

-   Compose a function `Gamma` which accepts one value, `z`.  `Gamma` should return the solution at that point according to the formula for $\Gamma$ above.

<!-- grade_id: gamma, points: 0, grade: false, solution: true -->
```python
# define your function here
def Gamma( z ):
    prod = 1
    for n in range(1,100):  # we use a large number but not actually infinity
        numer = _____
        denom = _____
        prod = prod * numer/denom
    term = 1 / z
    return term * prod
```

<!-- grade_id: gamma-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from math import gamma
from numpy import isclose
assert isclose(gamma(0.5),Gamma(0.5),rtol=1e-2)
assert isclose(gamma(1),Gamma(1))
assert isclose(gamma(1.5),Gamma(1.5),rtol=1e-2)
print( 'Success!' )
```

#### Napierian Logarithms

[John Napier invented the logarithm](https://en.wikipedia.org/wiki/Napierian_logarithm) as a calculation tool in 1614 with his publication of _Mirifici Logarithmorum Canonis Descriptio (Description of the Wonderful Rule of Logarithms)_.  Napier's logarithms were different from our current logarithms:  he preferred to not work with fractional or decimal parts, and so worked with large numbers instead.

The Napierian logarithm $\textrm{NapLog}$ may be defined as

$$
\textrm{NapLog}(x)
=
-10^{7} \textrm{log}_{e} \left( \frac{x}{10^{7}} \right)
\textrm{.}
$$

-   Compose a function `NapLog` which accepts one value, `x`.  `NapLog` should return the solution at that point according to the formula for $\textrm{NapLog}$ above.  You may import `math.log`.

<!-- grade_id: naplog, points: 0, grade: false, solution: true -->
```python
# define your function here
def NapLog( x ):
    return _____
```

<!-- grade_id: naplog-test1, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(NapLog(5000000),6931471)
assert isclose(NapLog(1e7),-0.0)
assert isclose(NapLog(3333333),10986123)
print( 'Success!' )
```