# Elementary Climate Calculations with SymPy

![](./img/lab05-header-bkgd.png)

❖ Objectives

-   Apply scientific modeling across a series of case studies.

### Foundations of Climate Modeling

Climate modeling seeks to describe the forward evolution of climate according to mass and energy flows and balances.  Climate can be construed to model any combination of temperature, precipitation, glacial mass, radiation, oceanography, and more.  Climate models vary from "zeroth dimensional" energy balances to complex multidimensional grids tracking mass and energy control volumes.

![](https://blogs.ei.columbia.edu/wp-content/uploads/2012/02/newglobe-300x240.jpg)

(Image courtesy Krajick2012.)

(We draw on [Wikipedia](https://en.wikipedia.org/wiki/Climate_model), [IPCC](https://web.archive.org/web/20030325024912/http://www.grida.no/climate/ipcc_tar/wg1/258.htm), and below-cited sources for elements of this lesson.)

### Zeroth-Dimensional Modeling

A zero-dimensional model considers the Earth as a single point subject to energy flux.  Balancing incoming and outgoing energy leads to this model,

$$
(1-a)s \pi r^2
= 
4 \pi r^2 \epsilon \sigma T^4
$$

where the left-hand side represents solar energy (incoming) and the right-hand side represents radiated black-body heat (outgoing) based on an "average" temperature of Earth.  Other parameters are:

- $s$, the solar constant, or the incoming solar radiation per unit area.  $s = 1367 \,\textrm{W}\cdot\textrm{m}^{-2}$.

- $a$, the Earth's average albedo (reflectivity).  $a = 0.3$.

- $r$, the Earth's radius. $r = 6.371 \times 10^{6} \,\textrm{m}$.

- $\sigma$, the Stefan-Boltzmann constant.  $\sigma = 5.67 \times 10^{-8} \,\textrm{J}\cdot\textrm{K}^{-4}\cdot\textrm{m}^{-2}\cdot\textrm{s}^{-1}$.

- $\epsilon$, the Earth's effective emissivity.  $\epsilon = 0.612$.

- Define the equation using SymPy and symbolic values for each parameter.

```python
import sympy
s,a,r,sigma,eps,T = sympy.S('s,a,r,sigma,epsilon,T')

lhs = (1-a)*s*sympy.pi*r**2
rhs = 4*sympy.pi*r**2*e*sigma*T**4
```

- Solve the equation for temperature (symbolically).  Since there are four answers, set `T_ans` equal to the positive real one only.

```python
T_ans = sympy.solve(lhs-rhs,T)[1]
print(T_ans)
```

- Solve the equation for temperature (numerically).  (If you have not seen how to do this yet in class, you can use `.subs` to substitute a value for a variable in SymPy.)

```python
T_final = float(T_ans.subs(s,1367).subs(a,0.3).subs(r,6.371e6).subs(sigma,5.67e-8).subs(eps,0.612))
print(T_final)
```

<!-- grade_id: calc_freq-test2, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(T_final,288.15)
print('Success!')
```

This widget allows you to manipulate the parameters and see the system response as a result of different albedos, emissivities, etc.

```python
from ipywidgets import interact
import matplotlib.pyplot as plt
%matplotlib inline
from math import ceil

x = []
y = []

def plot_temperature(s,a):
    r = 6.371e6
    sigma = 5.67e-8
    eps = 0.612
    T = 2**0.5*(-a*s/(eps*sigma) + s/(eps*sigma))**(1/4)/2
    x.append(len(x))
    y.append(T)
    plt.plot(x,y,'r.')
    plt.plot(len(x),T,'ro')
    plt.text(len(x),T,f'  {T-273:.1f} °C')
    plt.xlim((0,ceil(len(x)/100)*100))
    plt.ylim((0,400))
    plt.show()

interact(plot_temperature,s=(0,2500),a=(0,1,0.01))
```

### One-Dimensional Modeling

A one-dimensional climate model typically expands on the zeroth-dimensional radiative picture by including convection as well.  We will implement a model based on Mann's work (summarized in Miller2020a and Sokolik2007) which introduces a thickness to the atmosphere.

Our first model didn't separate the Earth from solar radiation, but modeled it as a single point.  We now introduce an "atmosphere," or at least a "distance" between the top of the atmosphere and the bottom (the Earth).  We therefore need to model three energy balances:

1.  The top of the atmosphere (with temperature $T_{e}$).
2.  The atmosphere itself.
3.  The Earth's surface (with temperature $T_{S}$.

![](https://www.e-education.psu.edu/meteo469/sites/www.e-education.psu.edu.meteo469/files/lesson05/Earth2Fig.gif)

(Image courtesy Miller2020a.)

We now have to consider the atmosphere with radiation and convection.  Radiation can be difficult to model because of the scattering and absorption of different wavelengths of energy.

![](http://www.severewx.com/Radiation/atmtrans.gif)

One assumption we can make when building this model is that the Earth acts as a "gray body," or partway between a [blackbody](https://en.wikipedia.org/wiki/Black_body) (perfect radiator) and "white body" (perfect reflector).  This makes things easier because then all of the radiation scatters the same way (and we don't have to model frequency as well).  Furthermore, we won't have any absorption in the atmosphere in our model.

At the top of the atmosphere, we have the following balance:

$$
\frac{s(1-a)}{4}
=
\sigma \epsilon T_{e}^{4} + (1-\epsilon)\sigma T_{S}^{4}
$$

which is a balance of incoming radiation on the left and outgoing radiation on the right.

The atmospheric layer has the following balance:

$$
\sigma \epsilon T_{S}^{4}
=
2\sigma\epsilon T_{e}^{4}
\textrm{.}
$$

The Earth's surface has the following balance:

$$
\frac{s(1-a)}{4} + \sigma\epsilon T_{e}^{4}
=
\sigma T_{S}^{4}
\textrm{.}
$$





- Solve for $T_{S}$ and $T_{e}$ using SympPy (and perhaps some algebra by hand) and symbolic values for each parameter.  You should be able to solve for $T_{S}$ without reference to $T_{e}$, and for $T_{e}$ as an equation in $T_{S}$.  Select positive real solutions whenever a choice arises.

<!-- grade_id: p2f, points: 0, grade: false, solution: true -->
```python
import sympy
s,a,sigma,eps,Te,Ts = sympy.S('s,a,sigma,epsilon,T_e,T_s')

toa_lhs = (1-a)*s/4
toa_rhs = sigma*eps*Te**4 + (1-eps)*sigma*Ts**4

atm_lhs = sigma*eps*Ts**4
atm_rhs = 2*sigma*eps*Te**4

sur_lhs = (1-a)*s/4 + sigma*eps*Te**4
sur_rhs = sigma*eps*Ts**4

Te_1 = sympy.solve(atm_lhs-atm_rhs,Te)
Ts_1 = sympy.solve((toa_lhs-toa_rhs).subs(Te,Te_1[1]),Ts)

print(Ts_1[3])
```

<!-- grade_id: p2f-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
test_Ts_1 = 215.4856
assert isclose(float(Ts_1[3].subs(a,0.78).subs(eps,0.77).subs(sigma,5.67e-8).subs(s,1367)),test_Ts_1)
print('Success!')
```

This widget allows you to manipulate the parameters and see the system response as a result of different albedos and emissivities, etc.  In particular, try these cases:

1.  $\epsilon = 0$ (no greenhouse effect)
2.  $\epsilon = 1$ (perfect greenhouse effect, fully absorptive atmosphere)
3.  $\epsilon = 0.77$ ("actual" greenhouse effect on Earth)
4.  Mercury:  $a = 0.06$
5.  Venus:  $a = 0.78$
6.  Earth:  $a = 0.78$
7.  Mars:  $a = 0.17$
8.  Jupiter:  $a = 0.45$

```python
from ipywidgets import interact
import matplotlib.pyplot as plt
%matplotlib inline
from math import ceil

x = []
y1 = []
y2 = []

def plot_temperature(s,a):
    r = 6.371e6
    sigma = 5.67e-8
    eps = 0.612
    Ts = ((a-1)*s/(2*eps*sigma - 4*sigma))**(1/4)
    Te = Ts/2**0.25
    x.append(len(x))
    y1.append(Ts)
    y2.append(Te)
    plt.plot(x,y1,'r.')
    plt.plot(x,y2,'b.')
    plt.plot(len(x),Ts,'ro')
    plt.plot(len(x),Te,'bo')
    plt.text(len(x),Ts,f'  {Ts-273:.1f} °C')
    plt.text(len(x),Te,f'  {Te-273:.1f} °C')
    plt.xlim((0,ceil(len(x)/100)*100))
    plt.ylim((0,400))
    plt.show()

interact(plot_temperature,s=(0,2500),a=(0,1,0.01))
```

### Three-Dimensional Global Climate Models

Full climate modeling most commonly slices the Earth's atmosphere and surface layers into many tiny prisms ("cells").  Each prism is modeled with a temperature, a flux rate, a heat capacity, etc.  This modeling involves many kinds of physics:

- fluid mechanics, to capture convection and turbulence within the atmosphere and the oceans
    - meteorology, to capture precipitation and atmospheric capacity
- heat transfer, to describe transfer between model cells
- radiative modeling, including varying rates of absorption by different gases and reflection by diverse surface types
- mass transfer, to describe water and gas transfer

![](https://www.gfdl.noaa.gov/wp-content/uploads/pix/model_development/climate_modeling/climatemodel.png)

(Image courtesy NOAA.)

### References

- [Krajick2012](https://blogs.ei.columbia.edu/2012/02/17/global-climate-modeling-for-the-masses-you-can-try-this-at-home/)

- [Manabe1964](https://journals.ametsoc.org/doi/abs/10.1175/1520-0469%281964%29021%3C0361%3ATEOTAW%3E2.0.CO%3B2)  S. Manabe, R. F. Strickler, 1964:  Thermal Equilibrium of the Atmosphere with a Convective Adjustment.  _Journal of the Atmospheric Sciences 24,_ 241–259.  doi:

- [Manabe1967](https://journals.ametsoc.org/doi/abs/10.1175/1520-0469(1967)024%3C0241%3ATEOTAW%3E2.0.CO%3B2)  S. Manabe, R. T. Wetherald, 1967:  Thermal Equilibrium of the Atmosphere with a Given Distribution of Relative Humidity.  _Journal of the Atmospheric Sciences 21,_ 361–385.  doi:10.1175/1520-0469(1964)021<0361:TEOTAW>2.0.CO;2


- [Miller2020a](https://www.e-education.psu.edu/meteo469/node/198)

- [Miller2020b](https://www.e-education.psu.edu/meteo469/node/212)

- [Sokolik2007](http://irina.eas.gatech.edu/EAS8803_FALL2007/), [Lecture 24](http://irina.eas.gatech.edu/EAS8803_FALL2007/Lecture24.pdf)

see also Wang1976 and Wang1980

- [Wang1976](https://science.sciencemag.org/content/194/4266/685).  W. C. Wang, Y. L. Yung, A. A. Lacis, T. Mo, J. E. Hansen, 1976: Greenhouse Effects due to Man-Made Perturbations of Trace Gases.  _Science 194_ (4266), 685–690.  doi:10.1126/science.194.4266.685

- [Wang1980](https://pubs.giss.nasa.gov/abs/wa03100m.html).  W. C. Wang, P.H. Stone, 1980: Effect of ice-albedo feedback on global sensitivity in a one-dimensional radiative-convective climate model. _J. Atmos. Sci., 37_, 545–552, doi:10.1175/1520-0469(1980)037<0545:EOIAFO>2.0.CO;2.-   Which language is the worst match, and its value of $f$?

```python

```