## `lab05`—Scientific Applications with `sympy`

![](./img/header.png)

Objectives

-   Plot expressions using SymPy.
-   Apply scientific modeling across a series of case studies.

### Kepler's Laws

In 1617, Johannes Kepler published _Epitome astronomiae Copernicanae_ (or _Epitome of Copernican Astronomy)_, which introduced his celebrated laws of motion.

[Kepler's three laws of motion](https://en.wikibooks.org/wiki/General_Astronomy/Kepler%27s_Laws) introduce a mathematical model of elliptical (rather than circular) orbits.  They are:

1.  The orbits of the planets are ellipses, with the sun at one focus.
2.  The line between a planet and the sun sweeps out equal areas in the plane of the planet's orbit over equal times.
3.  The time required for a planet to orbit the sun, called its period, is proportional to half the long axis of the ellipse raised to the 3/2 power.

In this portion of the lab, you will carry out calculations in support or demonstration of these laws.

##### Elliptical Orbits

An elliptical orbit has two bodies moving such that the lighter body orbits the heavier body as one focus of an ellipse.

![](https://upload.wikimedia.org/wikipedia/commons/c/c9/Kepler1.gif)

For real planets, the foci are quite close, as planetary orbits are only _slightly_ eccentric.  For comets and other bodies, the foci may lie far apart.

An equation describing the elliptical orbit may be written:

$$
\frac{x^2}{a^2} + \frac{y^2}{b^2}
=
1
$$

with eccentricity $e$ described by $b^2 = a^2\left( 1 - e^2 \right)$.  $a$ and $b$ are the _semimajor_ and _semiminor_ axes, respectively, of the ellipse.

```python
import sympy
x,y,a,b,e = sympy.S( 'x,y,a,b,e' )
```

In preparation for solving and plotting an orbit, we suppose that we are given the eccentricity $e$ and semimajor axis $a$, and need to obtain $b$ from that information.  First, solve for $b$ as `b_soln`:

```python
b_soln = ___

print( b_soln )
```

Next, `subs`titute in the values `e = 0.2056` and `a = 57909050000`, using SI units for convenience.  (These are the values for the planet Mercury ☿.)

```python
e_ = 0.2056
a_ = 57909050000
b_ = b_soln[ 1 ].___

print( 'e = %f; a = %f; b = %f' % ( e_,a_,b_ ) )
```

Now that we know $a$ and $b$, we can solve for the elliptical orbit.

Solving for $y$ in terms of $x$ will lead to multiple expressions since the ellipse is multivalued along the $x$ axis.

```python
y_ = sympy.solve( ___,y )

print( y_ )
```

We plot using a library we haven't seen yet, MatPlotLib.  This is actually the library behind SymPy's plotting ability, and we'll examine it during `num/plotting`.

A true circle is shown in red underlying the eccentric ellipse.

```python
# Plotting boilerplate---leave this as-is.
import matplotlib as mpl
import matplotlib.pyplot as plt
%matplotlib inline
import seaborn as sns
sns.set_context( 'talk',font_scale=1.5,rc={ 'lines.linewidth':5 } )
```

```python
# Plotting code.
plot1 = sympy.plotting.plot_parametric( ( x,y_[ 0 ] ),( x,-1e12,1e12 ),show=False )
plot2 = sympy.plotting.plot_parametric( ( x,y_[ 1 ] ),( x,-1e12,1e12 ),show=False )
plot_circle1 = sympy.plotting.plot_parametric( ( x,+(-x**2 + 3353458071902500000000)**0.5 ),( x,-1e12,1e12 ),show=False,line_color='red' )
plot_circle2 = sympy.plotting.plot_parametric( ( x,-(-x**2 + 3353458071902500000000)**0.5 ),( x,-1e12,1e12 ),show=False,line_color='red' )
plot_circle1.append( plot_circle2[ 0 ] )
plot_circle1.append( plot1[ 0 ] )
plot_circle1.append( plot2[ 0 ] )
mpl.rcParams[ 'figure.figsize' ] = 10,10   # make the figure square, necessary in Jupyter
plot_circle1.show()
```

##### Equal Area in Equal Time

Kepler's second law posits that the orbiting planet's velocity varies such that the same area is swept out during the same time.

![](https://upload.wikimedia.org/wikipedia/commons/7/75/Kepler2.gif)

![](https://upload.wikimedia.org/wikipedia/commons/thumb/6/69/Kepler-second-law.gif/220px-Kepler-second-law.gif)

As you can imagine, this becomes a somewhat complicated integral to write, depending on the term we solve for.  For convenience, we will switch to polar terms in our expression.

$$
\frac{dA}{dt}
=
\frac{1}{2}r^2 \frac{d\theta}{dt}
$$

and thus some calculus and algebra yield

$$
A(t)
=
\int_{\theta_0}^{\theta_1} d\theta\, \frac{r^2}{2}
=
\frac{1}{2} ab \left( \theta_1 - \theta_0 \right)
\text{.}
$$

The largest $t$ is the period of the orbit, and the planet sweeps out the area of the ellipse $\pi ab$ during that time.

$$
A(P)
=
\frac{1}{2} ab \left( 2\pi - 0 \right)
=
\pi ab
$$

Use SymPy to verify that the area traced from $\theta \in \left[ 0,\frac{\pi}{4} \right]$ is the same as that traced from $\theta \in \left[ \frac{3\pi}{8},\frac{5\pi}{8} \right]$.

<!-- grade_id: cell-2ada0376f491a468, points: 0, grade: false, solution: true -->
```python
area_1 = ___
area_2 = ___
```

<!-- grade_id: cell-0903030db1d816bc, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from math import isclose
assert ( area_1 == area_2 ) and ( area_1 == sympy.pi*a*b/8 )
print( 'Success!' )
```

##### Proportionality of Periods

Given the period of a planet or satellite, one can determine how wide the orbit is, and _vice versa_.

$$
P^2
\propto
a^3
$$

where $P$ is the period in years and $a$ is the length of the semimajor axis in astronomical units (not SI units as above).

Calculate the period of Mars ♂︎ given that its semimajor axis is $a = 1.87 \,\text{AU}$.

<!-- grade_id: cell-ab084336c256f58e, points: 0, grade: false, solution: true -->
```python
P = ___
```

<!-- grade_id: cell-1008d824d54d7721, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
test_P,test_a = sympy.S( 'P,a' )
from math import isclose
assert ( P == test_a.subs( test_a,1.87 ) ** ( 3 / 2 ) ) or ( isclose( P,1.87,rel_tol=1e-2 ) )
print( 'Success!' )
```

Calculate the semimajor axis of Jupiter ♃ given that its period is $11.86 \,\text{yr}$.

<!-- grade_id: cell-72b3e4058171eaaf, points: 0, grade: false, solution: true -->
```python
a = ___
```

<!-- grade_id: cell-839539944c4d5c32, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
test_P,test_a = sympy.S( 'P,a' )
from math import isclose
assert ( a == test_P.subs( test_P,11.86 ) ** ( 2 / 3 ) ) or ( isclose( P,5.2,rel_tol=1e-2 ) )
print( 'Success!' )
```

### Trajectories

A basic ballistic trajectory is subject only to the initial impetus.  More advanced trajectory calculations include thrust during flight and the action of external forces such as gravity.  Drag can be considered among the relevant forces as well.

In this section, we will calculate some trajectories by position and velocity.

##### Simple Trajectory

A simple trajectory describes an object moving only from the initial impulse, such as a rocket in deep space with the engines powered down.  We write the velocity as constant $v_0$, and integrate to obtain the position $x(t)$ as a function of $v_0$.

$$
v(t)
=
v_0
$$

$$
x(t)
=
\int_{\tau=t_0}^{t} d\tau\, v_0
$$

(We will use the Greek equivalent of Roman letters as the variables of integration since $t$ means the current time.)

```python
x,x0,v,v0,tau,t,tf = sympy.S( 'x,x0,v,v0,tau,t,tf' )
v = ___
x = sympy.integrate( v.subs( t,tau ),( tau,0,t ) )   # Note that tau not t is our variable of integration, so we subs it.
```

We plot velocity $v$ in red and position $x$ in blue.

```python
plot_v = sympy.plotting.plot( v.subs( v0,10 ),( t,0,10 ),line_color='red',show=False )
plot_x = sympy.plotting.plot( x.subs( v0,10 ),( t,0,10 ),show=False )
plot_v.append( plot_x[ 0 ] )
plot_v.show()
```

##### Powered Trajectory

That case was basically trivial, so let's try one subject to constant acceleration from a source like an [ion engine]().

$$
v(t)
=
a_0 t
$$

$$
x(t)
=
\int_{\tau=t_0}^{t} d\tau\, a_0 \tau
$$

```python
x,x0,v,a0,tau,t,tf = sympy.S( 'x,x0,v,a0,tau,t,tf' )
v = ___
x = ___  #should be able to write this the same as before
```

Once again, we plot velocity and position for this trajectory:

```python
plot_v = sympy.plotting.plot( v.subs( a0,10 ),( t,0,10 ),line_color='red',show=False )
plot_x = sympy.plotting.plot( x.subs( a0,10 ),( t,0,10 ),show=False )
plot_v.append( plot_x[ 0 ] )
plot_v.show()
```

##### External Forces

Next we consider a scenario with external forces including drag $a_D$ acting on the rocket.  In this case, we posit that acceleration due to drag will vary as a function of time (really, of position, but we're simplifying) as the rocket moves through a thinning atmosphere, $a_D = \exp\left( 2-t \right)$.  Since gravity is a constant acceleration and we are treating the rocket's thrust as constant, the behavior won't be fundamentally different.

$$
a(t)
=
a_0 + a_D(t)
=
a_0 + \exp\left( 2-t \right)
$$

$$
v(t)
=
\int_{\tau=t_0}^{t} d\tau\, a_0 + \exp\left( 10-t \right)
$$

$$
x(t)
=
\int_{\tau=t_0}^{t} d\tau\, v(\tau)
$$

```python
x,x0,v,a0,tau,t,tf = sympy.S( 'x,x0,v,a0,tau,t,tf' )
a = a0 - sympy.exp( 2-t )
v = ___
x = ___  #should be able to write this the same as before
```

Plotting the same:

```python
plot_v = sympy.plotting.plot( v.subs( a0,10 ),( t,0,10 ),line_color='red',show=False )
plot_x = sympy.plotting.plot( x.subs( a0,10 ),( t,0,10 ),show=False )
plot_v.append( plot_x[ 0 ] )
plot_v.show()
```

##### Odd Accelerations

An experimental rocket fuel has been found to burn with a strange sinusoidal acceleration, $a_F = 50 \sin( 10 t )$.  Subject to a standard acceleration $a_0 = -\frac{1}{2}$, will this fuel be able to free the rocket from the planet's gravitational field, or will the ride be too bumpy and ultimately fail?

To answer this question, you'll need to construct $a$ as a function of $a_0$ and $a_F$.

```python
x,x0,v,a0,tau,t,tf = sympy.S( 'x,x0,v,a0,tau,t,tf' )
a = 50*sympy.sin( 10*t ) + a0
v = ___
x = ___  #should be able to write this the same as before
```

Plot the resulting velocity and trajectory:

```python
plot_v = sympy.plotting.plot( v.subs( a0,-0.5 ),( t,0,10 ),line_color='red',show=False )
plot_x = sympy.plotting.plot( x.subs( a0,-0.5 ),( t,0,10 ),show=False )
plot_v.append( plot_x[ 0 ] )
plot_v.show()
```

A different rocket fuel has been found to burn with the acceleration, $a_F = 2 \log( 10 t )$.  Subject to a standard acceleration $a_0 = -\frac{1}{2}$, check this fuel on the same planet to see if the rocket can escape the gravitational field.

To answer this question, you'll need to construct $a$ as a function of $a_0$ and $a_F$.

```python
x,x0,v,a0,tau,t,tf = sympy.S( 'x,x0,v,a0,tau,t,tf' )
a = 2*sympy.log( 10 * t ) + a0
v = ___
x = ___  #should be able to write this the same as before
```

```python
plot_v = sympy.plotting.plot( v.subs( a0,-0.5 ),( t,0,10 ),line_color='red',show=False )
plot_x = sympy.plotting.plot( x.subs( a0,-0.5 ),( t,0,10 ),show=False )
plot_v.append( plot_x[ 0 ] )
plot_v.show()
```

### Foundations of Climate Modeling

Climate modeling seeks to describe the forward evolution of climate according to mass and energy flows and balances.  Climate can be construed to model any combination of temperature, precipitation, glacial mass, radiation, oceanography, and more.  Climate models vary from "zeroth dimensional" energy balances to complex multidimensional grids tracking mass and energy control volumes.

![](https://blogs.ei.columbia.edu/wp-content/uploads/2012/02/newglobe-300x240.jpg)

(Image courtesy Krajick2012.)

(We draw on [Wikipedia](https://en.wikipedia.org/wiki/Climate_model), [IPCC](https://web.archive.org/web/20030325024912/http://www.grida.no/climate/ipcc_tar/wg1/258.htm), and below-cited sources for elements of this lesson.)

### Zeroth-Dimensional Modeling

A zero-dimensional model considers the Earth as a single point subject to energy flux.  Balancing incoming and outgoing energy leads to this model,

$$
(1-a)s \pi r^2
= 
4 \pi r^2 \epsilon \sigma T^4
$$

where the left-hand side represents solar energy (incoming) and the right-hand side represents radiated black-body heat (outgoing) based on an "average" temperature of Earth.  Other parameters are:

- $s$, the solar constant, or the incoming solar radiation per unit area.  $s = 1367 \,\textrm{W}\cdot\textrm{m}^{-2}$.

- $a$, the Earth's average albedo (reflectivity).  $a = 0.3$.

- $r$, the Earth's radius. $r = 6.371 \times 10^{6} \,\textrm{m}$.

- $\sigma$, the Stefan-Boltzmann constant.  $\sigma = 5.67 \times 10^{-8} \,\textrm{J}\cdot\textrm{K}^{-4}\cdot\textrm{m}^{-2}\cdot\textrm{s}^{-1}$.

- $\epsilon$, the Earth's effective emissivity.  $\epsilon = 0.612$.

- Define the equation using SymPy and symbolic values for each parameter.

```python
import sympy
s,a,r,sigma,eps,T = sympy.S('s,a,r,sigma,epsilon,T')

lhs = (1-a)*s*sympy.pi*r**2
rhs = 4*sympy.pi*r**2*eps*sigma*T**4
```

- Solve the equation for temperature (symbolically).  Since there are four answers, set `T_ans` equal to the positive real one only.

```python
T_ans = sympy.solve(lhs-rhs,T)[1]
print(T_ans)
```

Output:
```
sqrt(2)*(-a*s/(epsilon*sigma) + s/(epsilon*sigma))**(1/4)/2
```

- Solve the equation for temperature (numerically).  (If you have not seen how to do this yet in class, you can use `.subs` to substitute a value for a variable in SymPy.  Your TA or CA can clarify.)

```python
T_final = float(T_ans.subs(s,1367).subs(a,0.3).subs(r,6.371e6).subs(sigma,5.67e-8).subs(eps,0.612))
print(T_final)
```

Output:
```
288.14959967471793
```

<!-- grade_id: calc_freq-test2, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(T_final,288.15)
print('Success!')
```

Output:
```
Success!
```

This widget allows you to manipulate the parameters and see the system response as a result of different albedos, emissivities, etc.

```python
from ipywidgets import interact
import matplotlib.pyplot as plt
%matplotlib inline
from math import ceil

x = []
y = []

def plot_temperature(s,a):
    r = 6.371e6
    sigma = 5.67e-8
    eps = 0.612
    T = 2**0.5*(-a*s/(eps*sigma) + s/(eps*sigma))**(1/4)/2
    x.append(len(x))
    y.append(T)
    plt.plot(x,y,'r.')
    plt.plot(len(x),T,'ro')
    plt.text(len(x),T,f'  {T-273:.1f} °C')
    plt.xlim((0,ceil(len(x)/100)*100))
    plt.ylim((0,400))
    plt.show()

interact(plot_temperature,s=(0,2500),a=(0,1,0.01))
```

Output:
```
interactive(children=(IntSlider(value=1250, description='s', max=2500), FloatSlider(value=0.0, description='a'…
<function __main__.plot_temperature(s, a)>
```

### One-Dimensional Modeling

A one-dimensional climate model typically expands on the zeroth-dimensional radiative picture by including convection as well.  We will implement a model based on Mann's work (summarized in Miller2020a and Sokolik2007) which introduces a thickness to the atmosphere.

Our first model didn't separate the Earth from solar radiation, but modeled it as a single point.  We now introduce an "atmosphere," or at least a "distance" between the top of the atmosphere and the bottom (the Earth).  We therefore need to model three energy balances:

1.  The top of the atmosphere (with temperature $T_{e}$).
2.  The atmosphere itself.
3.  The Earth's surface (with temperature $T_{S}$.

![](https://www.e-education.psu.edu/meteo469/sites/www.e-education.psu.edu.meteo469/files/lesson05/Earth2Fig.gif)

(Image courtesy Miller2020a.)

We now have to consider the atmosphere with radiation and convection.  Radiation can be difficult to model because of the scattering and absorption of different wavelengths of energy.

![](http://www.severewx.com/Radiation/atmtrans.gif)

One assumption we can make when building this model is that the Earth acts as a "gray body," or partway between a [blackbody](https://en.wikipedia.org/wiki/Black_body) (perfect radiator) and "white body" (perfect reflector).  This makes things easier because then all of the radiation scatters the same way (and we don't have to model frequency as well).  Furthermore, we won't have any absorption in the atmosphere in our model.

At the top of the atmosphere, we have the following balance:

$$
\frac{s(1-a)}{4}
=
\sigma \epsilon T_{e}^{4} + (1-\epsilon)\sigma T_{S}^{4}
$$

which is a balance of incoming radiation on the left and outgoing radiation on the right.

The atmospheric layer has the following balance:

$$
\sigma \epsilon T_{S}^{4}
=
2\sigma\epsilon T_{e}^{4}
\textrm{.}
$$

The Earth's surface has the following balance:

$$
\frac{s(1-a)}{4} + \sigma\epsilon T_{e}^{4}
=
\sigma T_{S}^{4}
\textrm{.}
$$





- Solve for $T_{S}$ and $T_{e}$ using SympPy (and perhaps some algebra by hand) and symbolic values for each parameter.  You should be able to solve for $T_{S}$ without reference to $T_{e}$, and for $T_{e}$ as an equation in $T_{S}$.  Select positive real solutions whenever a choice arises.

<!-- grade_id: p2f, points: 0, grade: false, solution: true -->
```python
import sympy
s,a,sigma,eps,Te,Ts = sympy.S('s,a,sigma,epsilon,T_e,T_s')

toa_lhs = (1-a)*s/4
toa_rhs = sigma*eps*Te**4 + (1-eps)*sigma*Ts**4

atm_lhs = sigma*eps*Ts**4
atm_rhs = 2*sigma*eps*Te**4

sur_lhs = (1-a)*s/4 + sigma*eps*Te**4
sur_rhs = sigma*eps*Ts**4

Te_1 = sympy.solve(atm_lhs-atm_rhs,Te)
Ts_1 = sympy.solve((toa_lhs-toa_rhs).subs(Te,Te_1[1]),Ts)

print(Ts_1[3])
```

Output:
```
(a*s/(2*epsilon*sigma - 4*sigma) - s/(2*epsilon*sigma - 4*sigma))**(1/4)
```

<!-- grade_id: p2f-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
print(Ts_1[3].subs(eps,0.77).subs(sigma,5.67e-8))
assert isclose(float(Ts_1[3].subs(eps,0.77).subs(sigma,5.67e-8).subs(a,0.78).subs(s,)),288.15)
print('Success!')
```

Output:
```
(-7169383.8631508*a*s + 7169383.8631508*s)**(1/4)
<ipython-input-13-5f4a8cada23e> in <module>
      2 from numpy import isclose
      3 print(Ts_1[3].subs(eps,0.77).subs(sigma,5.67e-8))
----> 4 assert isclose(float(Ts_1[3].subs(eps,0.77).subs(sigma,5.67e-8)),288.15)
      5 print('Success!')

~/Code/anaconda3/lib/python3.7/site-packages/sympy/core/expr.py in __float__(self)
    278         if result.is_number and result.as_real_imag()[1]:
    279             raise TypeError("can't convert complex to float")
--> 280         raise TypeError("can't convert expression to float")
    281 
    282     def __complex__(self):

TypeError: can't convert expression to float
```

This widget allows you to manipulate the parameters and see the system response as a result of different albedos and emissivities, etc.  In particular, try these cases:

1.  $\epsilon = 0$ (no greenhouse effect)
2.  $\epsilon = 1$ (perfect greenhouse effect, fully absorptive atmosphere)
3.  $\epsilon = 0.77$ ("actual" greenhouse effect on Earth)
4.  Mercury:  $a = 0.06$
5.  Venus:  $a = 0.78$
6.  Earth:  $a = 0.78$
7.  Mars:  $a = 0.17$
8.  Jupiter:  $a = 0.45$

```python
from ipywidgets import interact
import matplotlib.pyplot as plt
%matplotlib inline
from math import ceil

x = []
y1 = []
y2 = []

def plot_temperature(s,a):
    r = 6.371e6
    sigma = 5.67e-8
    eps = 0.612
    Ts = ((a-1)*s/(2*eps*sigma - 4*sigma))**(1/4)
    Te = Ts/2**0.25
    x.append(len(x))
    y1.append(Ts)
    y2.append(Te)
    plt.plot(x,y1,'r.')
    plt.plot(x,y2,'b.')
    plt.plot(len(x),Ts,'ro')
    plt.plot(len(x),Te,'bo')
    plt.text(len(x),Ts,f'  {Ts-273:.1f} °C')
    plt.text(len(x),Te,f'  {Te-273:.1f} °C')
    plt.xlim((0,ceil(len(x)/100)*100))
    plt.ylim((0,400))
    plt.show()

interact(plot_temperature,s=(0,2500),a=(0,1,0.01))
```